// exercises/src/conversions.rs
// Chapter 23: 类型转换 — rustlings fork
// 深做章节

// 每题含 `// I AM NOT DONE` 注释，运行前删除即可

#[test]
#[ignore]
    // I AM NOT DONE
fn conversions1() {
    // I AM NOT DONE
    todo!();
}

// Exercise conversions2
// 为自定义类型实现 From<&str> 和 From<String>。
#[test]
#[ignore]
    // I AM NOT DONE
fn conversions2() {












    // I AM NOT DONE
    todo!();
}

// Exercise conversions3
// 实现 TryFrom<i64>，拒绝负数并返回自定义错误。
#[test]
#[ignore]
    // I AM NOT DONE
fn conversions3() {
























    // I AM NOT DONE
    todo!();
}
