// exercises/src/enums.rs
// Chapter 08: Enums — rustlings fork
// 快进章节

// 枚举是 Rust 中表达"多选一"数据的方式。本章练习枚举定义、模式匹配和 match。

#[test]
#[ignore]
fn enums1() {
    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn enums2() {















    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn enums3() {









    // I AM NOT DONE
    todo!();
}
