// exercises/src/functions.rs
// Chapter 02: functions — rustlings fork
// 快进章节（预计 1h 内完成）

// 每题含 `// I AM NOT DONE` 注释，运行前删除即可

#[test]
#[ignore]
fn functions1() {
    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn functions2() {







    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn functions3() {






    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn functions4() {





    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn functions5() {






    // I AM NOT DONE
    todo!();
}
