// exercises/src/hashmaps.rs
// Chapter 11: HashMap — rustlings fork
// 深做章节

// 每题含 `// I AM NOT DONE` 注释，运行前删除即可

#![allow(dead_code, unused_imports)]

use std::collections::HashMap;

#[test]
#[ignore]
    // I AM NOT DONE
fn hashmaps1() {
    // I AM NOT DONE
    todo!();
}

// Exercise hashmaps2
// 使用 entry 与 or_insert 统计词频，忽略大小写。
#[test]
#[ignore]
    // I AM NOT DONE
fn hashmaps2() {






    // I AM NOT DONE
    todo!();
}

// Exercise hashmaps3
// 使用 entry API 在 key 不存在时插入默认值。
#[test]
#[ignore]
    // I AM NOT DONE
fn hashmaps3() {









    // I AM NOT DONE
    todo!();
}
