// exercises/src/if_.rs
// Chapter 03: if — rustlings fork
// 快进章节（预计 1h 内完成）

// 每题含 `// I AM NOT DONE` 注释，运行前删除即可

#[test]
#[ignore]
fn if1() {
    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn if2() {




    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn if3() {




    // I AM NOT DONE
    todo!();
}
