// exercises/src/intro.rs
// Chapter 00: intro — rustlings fork
// 快进章节（预计 1h 内完成）

// 每题含 `// I AM NOT DONE` 注释，运行前删除即可

#[test]
#[ignore]
fn intro1() {
    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn intro2() {


    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn intro3() {



    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn intro4() {


    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn intro5() {






    // I AM NOT DONE
    todo!();
}
