// exercises/src/iterators.rs
// Chapter 18: 迭代器 — rustlings fork
// 深做章节

// 每题含 `// I AM NOT DONE` 注释，运行前删除即可

#[test]
#[ignore]
    // I AM NOT DONE
fn iterators1() {
    // I AM NOT DONE
    todo!();
}

// Exercise iterators2
// 使用 filter 与 collect 保留 1 到 6 中的偶数。
#[test]
#[ignore]
    // I AM NOT DONE
fn iterators2() {
    // I AM NOT DONE
    todo!();
}

// Exercise iterators3
// 使用 find 与 copied 找到第一个大于 10 的数。
#[test]
#[ignore]
    // I AM NOT DONE
fn iterators3() {
    // I AM NOT DONE
    todo!();
}

// Exercise iterators4
// 使用 fold 累计字符串长度。
#[test]
#[ignore]
    // I AM NOT DONE
fn iterators4() {

    // I AM NOT DONE
    todo!();
}

// Exercise iterators5
// 使用 collect 把 Vec<Result<i32, &str>> 中的 Ok 值收集成 Result<Vec<i32>, &str>。
#[test]
#[ignore]
    // I AM NOT DONE
fn iterators5() {

    // I AM NOT DONE
    todo!();
}
