// exercises/src/lifetimes.rs
// Chapter 16: Lifetimes — rustlings fork
// 深做章节

// 生命周期标注帮助编译器验证引用有效性。本章覆盖函数、结构体和
// impl 块上的生命周期标注。

#[test]
#[ignore]
fn lifetimes1() {
    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn lifetimes2() {









    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn lifetimes3() {
















    // I AM NOT DONE
    todo!();
}
