// exercises/src/macros.rs
// Chapter 21: 宏 — rustlings fork
// 深做章节

// 每题含 `// I AM NOT DONE` 注释，运行前删除即可

#[test]
#[ignore]
    // I AM NOT DONE
fn macros1() {
    // I AM NOT DONE
    todo!();
}

// Exercise macros2
// 使用 format! 宏构造一个字符串。
#[test]
#[ignore]
    // I AM NOT DONE
fn macros2() {

    // I AM NOT DONE
    todo!();
}

// Exercise macros3
// 定义并使用一个简单的 macro_rules! 宏，计算参数个数。
#[test]
#[ignore]
    // I AM NOT DONE
fn macros3() {




    // I AM NOT DONE
    todo!();
}

// Exercise macros4
// 定义一个 macro_rules! 宏，实现类似 vec![x; n] 的重复模式。
#[test]
#[ignore]
    // I AM NOT DONE
fn macros4() {













    // I AM NOT DONE
    todo!();
}
