// exercises/src/move_semantics.rs
// Chapter 06: move_semantics — rustlings fork
// 快进章节（预计 1h 内完成）

// 每题含 `// I AM NOT DONE` 注释，运行前删除即可

#[test]
#[ignore]
fn move_semantics1() {
    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn move_semantics2() {









    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn move_semantics3() {










    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn move_semantics4() {








    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn move_semantics5() {






    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn move_semantics6() {





    // I AM NOT DONE
    todo!();
}
