// exercises/src/options.rs
// Chapter 12: Option & Result — rustlings fork
// 深做章节

// 每题含 `// I AM NOT DONE` 注释，运行前删除即可

#[test]
#[ignore]
    // I AM NOT DONE
fn options1() {
    // I AM NOT DONE
    todo!();
}

// Exercise options2
// 从字符串解析 u32，失败时返回默认值 0（不使用 unwrap）。
#[test]
#[ignore]
    // I AM NOT DONE
fn options2() {







    // I AM NOT DONE
    todo!();
}

// Exercise options3
// 使用 Option 的组合方法，仅保留能被 3 整除的数，并将其转换为描述字符串。
#[test]
#[ignore]
    // I AM NOT DONE
fn options3() {








    // I AM NOT DONE
    todo!();
}
