// exercises/src/primitive_types.rs
// Chapter 04: primitive_types — rustlings fork
// 快进章节（预计 1h 内完成）

// 每题含 `// I AM NOT DONE` 注释，运行前删除即可

#[test]
#[ignore]
fn primitive_types1() {
    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn primitive_types2() {





    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn primitive_types3() {




    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn primitive_types4() {




    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn primitive_types5() {



    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn primitive_types6() {





    // I AM NOT DONE
    todo!();
}
