// exercises/src/smart_pointers.rs
// Chapter 19: 智能指针 — rustlings fork
// 深做章节

// 每题含 `// I AM NOT DONE` 注释，运行前删除即可

#[test]
#[ignore]
    // I AM NOT DONE
fn smart_pointers1() {
    // I AM NOT DONE
    todo!();
}

// Exercise smart_pointers2
// 使用 Rc 共享数据，并验证其强引用计数。
#[test]
#[ignore]
    // I AM NOT DONE
fn smart_pointers2() {
















    // I AM NOT DONE
    todo!();
}

// Exercise smart_pointers3
// 使用 RefCell 在不可变引用内部修改数据。
#[test]
#[ignore]
    // I AM NOT DONE
fn smart_pointers3() {










    // I AM NOT DONE
    todo!();
}
