// exercises/src/structs.rs
// Chapter 07: structs — rustlings fork
// 快进章节（预计 1h 内完成）

// 每题含 `// I AM NOT DONE` 注释，运行前删除即可

#[test]
#[ignore]
fn structs1() {
    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn structs2() {














    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn structs3() {
















    // I AM NOT DONE
    todo!();
}
