// exercises/src/tests.rs
// Chapter 17: 测试 — rustlings fork
// 深做章节

// 每题含 `// I AM NOT DONE` 注释，运行前删除即可

#[test]
#[ignore]
    // I AM NOT DONE
fn tests1() {
    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
    // I AM NOT DONE
fn tests2() {

    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
#[should_panic(expected = "division by zero")]
    // I AM NOT DONE
fn tests3() {

    // I AM NOT DONE
    todo!();
}
