// exercises/src/threads.rs
// Chapter 20: 线程 — rustlings fork
// 深做章节

// 每题含 `// I AM NOT DONE` 注释，运行前删除即可

#![allow(dead_code, unused_imports)]

use std::sync::{Arc, Mutex};
use std::thread;

#[test]
#[ignore]
    // I AM NOT DONE
fn threads1() {
    // I AM NOT DONE
    todo!();
}

// Exercise threads2
// 使用 move 闭包把数据传入线程。
#[test]
#[ignore]
    // I AM NOT DONE
fn threads2() {



    // I AM NOT DONE
    todo!();
}

// Exercise threads3
// 使用 Arc<Mutex<T>> 让多个线程安全地累加计数器。
#[test]
#[ignore]
    // I AM NOT DONE
fn threads3() {


    // I AM NOT DONE
    todo!();
}
