// exercises/src/traits.rs
// Chapter 15: Traits — rustlings fork
// 深做章节

// 每题含 `// I AM NOT DONE` 注释，运行前删除即可

#[test]
#[ignore]
    // I AM NOT DONE
fn traits1() {
    // I AM NOT DONE
    todo!();
}

// Exercise traits2
// 使用 #[derive] 让 Point 支持 Debug、PartialEq 和 Clone。
#[test]
#[ignore]
    // I AM NOT DONE
fn traits2() {











    // I AM NOT DONE
    todo!();
}

// Exercise traits3
// 实现 Summary trait，并为 Tweet 提供 @username: content 格式。
#[test]
#[ignore]
    // I AM NOT DONE
fn traits3() {








    // I AM NOT DONE
    todo!();
}

// Exercise traits4
// 使用泛型 trait bound 调用 greet。
#[test]
#[ignore]
    // I AM NOT DONE
fn traits4() {






































    // I AM NOT DONE
    todo!();
}

// Exercise traits5
// 实现 Drop trait，变量离开作用域时增加 DROPPED_COUNT。
#[test]
#[ignore]
    // I AM NOT DONE
fn traits5() {



















    // I AM NOT DONE
    todo!();
}
