// exercises/src/variables.rs
// Chapter 01: variables — rustlings fork
// 快进章节（预计 1h 内完成）

// 每题含 `// I AM NOT DONE` 注释，运行前删除即可

#[test]
#[ignore]
fn variables1() {
    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn variables2() {



    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn variables3() {


    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn variables4() {


    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn variables5() {




    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn variables6() {



    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn variables7() {




    // I AM NOT DONE
    todo!();
}

#[test]
#[ignore]
fn variables8() {



    // I AM NOT DONE
    todo!();
}
